# Assignment 01

for the course Introduction to Data Science and Programming (Autumn 2025) at ITU

**Due: September 11, 2025, at 23:59** (hand in through [learnIT](https://learnit.itu.dk/course/view.php?id=3024641)!)

Before you turn this assignment in, make sure everything runs as expected. (Go to `Restart and run all` - should not throw any assertion errors.) 

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and ITU email address below:

```python
NAME = ""
ITU_MAIL = "" # xxxx@itu.dk
```

***
**Mandatory Task 1:** Write the function `sum_list` that takes a list as input and returns the sum of all elements in the list as ouput. Do **not** use Python's built-in `sum()` function for this task; the task is to write this function from scratch. It is fine if your function throws an error if anything else than a list of integers is entered as input.

<!-- grade_id: cell-c3b8fc79af83886c, points: 0, grade: false, solution: true -->
```python
# Solve Task 1
def sum_list(my_list):
    '''
    your docstring here
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

Make sure your function passes all the tests: running the cell below should not throw any errors! If it does, it means you need to adjust your function. You can run this cell as often as you need to, to make sure your function works.

<!-- grade_id: cell-35af3afdbf1678c3, points: 1, grade: true, solution: false -->
```python
# Check if you solved Task 1 correctly by running this cell
assert sum_list([1,2,3])==6
assert sum_list([9])==9
assert sum_list([0,0,-1,-10])==-11

```

***
**Mandatory Task 2:** Write the function `is_odd` that takes an integer as input and returns either `True` (if the input integer is odd), `False` (if the input integer is even). It is fine if your function throws an error if anything else than an integer is entered as input.

<!-- grade_id: cell-a3ec09c0ee06a3e1, points: 0, grade: false, solution: true -->
```python
# Solve Task 2
def is_even(n):
    '''
    your docstring here
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7905e96f23995be4, points: 1, grade: true, solution: false -->
```python
# Check if you solved Task 2 correctly

# we correctly detect odd numbers
assert is_odd(1)==True
assert is_odd(7659)==True
assert is_odd(-3)==True

# we correctly detect even numbers
assert is_odd(2)==False
assert is_odd(16)==False
assert is_odd(0)==False
```

***
**Mandatory Task 3:** Write the function `compute_factorial` that takes an integer $n>=1$ as input and returns the factorial $n!$ as output. Do **not** use any built-in functions or modules in Python for this task. **Hint:** use a `for` loop with a `range()` to compute the factorial. 

<!-- grade_id: cell-e8db086fc0ce7341, points: 0, grade: false, solution: true -->
```python
# Solve Task 3
def compute_factorial(n):
    '''
    your docstring here
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-35c12a7e2e2818f4, points: 1, grade: true, solution: false -->
```python
# Check if you solved Task 3 correctly

# Computing factorials correctly
assert compute_factorial(1)==1
assert compute_factorial(3)==6
assert compute_factorial(10)==3628800
```

***
**Mandatory Task 4** Write a function `find_largest_factorial_below` that takes a positive number $n$ (`int` or `float`) as input and returns the **largest possible** number $a$ for which $a! \leq n$. In other words, this function returns the number $a$ whose factorial $a!$ gets as close as possible to $n$ without growing larger than $n$. **Hint:** Use a `while` loop with within your function.

<!-- grade_id: cell-8c5e7d0a48ac173b, points: 0, grade: false, solution: true -->
```python
# Solve Task 4
def find_largest_factorial_below(n):
    '''
    your docstring here
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d583b2ef04bcd208, points: 1, grade: true, solution: false -->
```python
# Check if you solved Task 4 correctly

# For numbers between 0 and 2, it should return 1
assert find_largest_factorial_below(1.5)==1

# For numbers between 6 and 24, it should return 3
assert find_largest_factorial_below(17)==3

# For 24, it should return 4
assert find_largest_factorial_below(24)==4
```

***

**Mandatory Task 5** Write a function `count_unique_characters` that
* takes as input a string,
* returns as output a dictionary of key-value pairs where keys are unique characters of the input string, and values are counts (number of times this character appears in the input string). 

For example, for the input string `"hello hello!"` the function should return the following dictionary:
```python
{
    "h": 2,
    "e": 2,
    "l": 4,
    "o": 2,
    " ": 1,
    "!": 1
}
``` 

```python
# Solve Task 5
def find_unique_characters(mystring):
    '''
    your docstring here
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Check if you solved Task 5 correctly

assert find_unique_characters("hello hello!") == {
    "h": 2,
    "e": 2,
    "l": 4,
    "o": 2,
    " ": 1,
    "!": 1
}

assert find_unique_characters("1 stitch in time...") == {
    'i': 3,
    's': 1,
    'e': 1,
    'n': 1,
    '.': 3,
    '1': 1,
    't': 3,
    ' ': 3,
    'c': 1,
    'h': 1,
    'm': 1
}
```

***

**Mandatory Task 6:** **In your own words**, explain line by line what happens in the code below.

```python
def should_i_bike(rain):
    if rain:
        return False
    else:
        return True

def should_i_bike_fast(late):
    if late:
        return True
    else:
        return False

rain_today = False
late_today = True

if should_i_bike(rain_today):
    if not should_i_bike_fast(late_today):
        print("bike peacefully")    
    else:
        print("bike fast")
else:
    print("take the metro")
```

<!-- grade_id: cell-7e11d193408739d6, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE